```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5

## Due Date: Wednesday, June 18th by 11:59 PM MT on Gradescope

**Assignment Learning Objectives:**

 - Determine whether two events are independent
 - Define random variables.  Explain the difference between random variables and events
 - Explain what a probability mass function (PMF) is and use it to calculate probabilities of discrete random variables.
 - Use tables, histograms and/or closed-form functions to represent PMFs
 - Determine whether 2 discrete RV are independent using the mathematical definition
 - Calculate and interpret the expected value of a random variable
 - Calculate and interpret the variance and standard deviation of a random variable


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Entering a problem into ChatGPT and then submitting the solution or a modification of the solution.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Discord!


#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 5 | autograded |
Question 2 | 13 | autograded & manual |
Question 3 | 15 | autograded & manual |
Question 4 | 10 | manual |
Question 5 | 7 | autograded & manual|
|Total | 50


## Assigned Readings for this HW:
- [CS109 Course Reader: Independence](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/independence/)
- [CS109 Course Reader: Independence in Variables](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/independent_vars/)
- [CS109 Course Reader: Random Variables](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/rvs/)
- [CS109 Course Reader: Probability Mass Functions (PMF)](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/pmf/)
- [CS109 Course Reader: Many Head Flips](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/many_flips/)
- [CS109 Course Reader: Expectation](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/expectation/)
- [CS109 Course Reader: Variance](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/variance/)
- [CS109 Course Reader: Binomial Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part2/binomial/)

<a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) | [Problem 5](#p4) | 
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



```python
# Run this cell to start

import pandas as pd
import numpy as np
import hashlib
import matplotlib.pylab as plt
%matplotlib inline


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

# Create a Generator instance:
rng = np.random.default_rng()
```

***

<a id='p1'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1: Independent Events - Video Assignment (5 pts) ## 

The word "independent" in mathematics has a very specific definition (and it turns out it doesn't always correspond to our common use of the English word 'independent').  In this question, we're going to dive into the mathematical definition of what it means for events to be independent or dependent.  

#### Complete this [Lesson 9: Independent Events Video Assignment (click this link)](https://canvas.colorado.edu/courses/120673/assignments/2371675) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).

<a id='p2'></a>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2 (13 points)
    
Suppose you roll two fair, 6-sided dice. 

Let $D_1$ be the random variable that denotes the outcome of the first die.  

Let $D_2$ be the random variable that denotes the outcome of the second die. 

Let $S$ be the random variable that denotes the sum of the 2 rolls.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2a (1 point)

What is the support of $D_1$?  (that is, what is the set of numbers that are possible values of this random variable)? Give your answer as a python **list**.

```python
d_sup = ...
```

```python
grader.check("2a")
```

Next we'll make a Pandas series that represents the distribution (aka the probability mass function) of $D_1$.

To create a series from scratch in Pandas, with a specific index, you can use the following syntax:

```python
s = pd.Series([-1, 10, 2], index = ["a", "b", "c"])
s
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2b (1 point)

Create a Pandas **series** that gives the the Probability Mass Function (PMF) of $D_1$. That is, the index of the series is the support of $D_1$ and the values are the probabilities.

```python
d1_pmf = ...
d1_pmf
```

```python
grader.check("2b")
```

**Plotting PMFs**

You can plot a histogram of the PMF of $D_1$ using the following syntax: 

**NOTE:** 
 The parameter rot=0 just rotates the labels on the x-axis to be horizontal. 

```python
d1_pmf.plot.bar(rot=0,width=1, ec='white')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2c (1 point)
    
What is the support of $S$?  (that is, what is the set of numbers that are possible values of this random variable)? Give your answer as a python **list**.    

```python
s_sup = ...
s_sup
```

```python
grader.check("2c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2d (1 point)
    
Create a **Pandas series** that gives the the PMF of $S$.    That is the index of the series is the support and the values are the probabilities.

```python
s_pmf = ...
s_pmf
```

```python
grader.check("2d")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2e (1 point)
    
Plot a histogram of the PMF of $S$.
(Use the same syntax showed above so that the bin widths are equal to 1, there's a white lines between each value and we can interpret area in the histogram as probabilities. Be sure to label the title and axes.)

```python
...
```

<!-- END QUESTION -->

**Independence**

Independence is a big deal in machine learning and probabilistic modeling. Knowing the “joint”
probability of many events (the probability of the “and” of the events) requires exponential amounts
of data. By making independence (and conditional independence) claims, computers can essentially
decompose how to calculate the joint probability, making it faster to compute, and requiring less
data to learn probabilities.  In this next part of the problem we'll practice with some basic concepts regarding independence.  

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2f (2 points)

Are events $D_1 =1$ and $D_2 = 1$ independent or dependent? Justify your answer using the mathematical definition of independence.

Write up your full solution in the box below using LaTeX (not code).  Show all steps fully justifying your answers.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2g (2 points)
    
Are events $D_1 = 3$ and $S=7$ independent or dependent?  Justify your answer using the mathematical definition of independence.

Write up your full solution in the box below using LaTeX (not code).  Show all steps fully justifying your answers.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2h (2 points)
    
Are events $D_1 = 1$ and $S=6$ independent or dependent?  Justify your answer using the mathematical definition of independence.

Write up your full solution in the box below using LaTeX (not code).  Show all steps fully justifying your answers.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2i (2 points)
    
Are the **random variables** $D_1$ and $S$ independent or dependent?  Justify your answer using the definition of independent random variables.

Write up your full solution in the box below using LaTeX (not code).  Show all steps fully justifying your answers.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<a id='p3'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3 (16 points)



A roulette wheel in the U.S. contains 38 equally sized spaces. The wheel is spun and a ball randomly lands in one of these spaces. Two spaces are green and have numbers 0 and 00 on them. The other spaces are numbered from 1 to 36. Half of these remaining spaces are red and half of them are black. Different wagers can be made on where the ball will end up landing. 

You can place bets on various number/color combinations and each type of bet pays-out at a different rate.  

We define **net winnings** to be the amount you win minus the amount you paid to play.

For example: 

- If you bet $\$1$ on red (or black) and win, then your **net winnings** is $\$1$ (i.e. you get your original dollar back, plus $\$1$).
- If you bet $\$1$ on red (or black) and lose, then your **net winnings** is -$\$1$ (i.e. you lose your original dollar).  
- If you bet $\$1$ any particular number and win, then your **net winnings** is $\$35$ (i.e you get your original dollar back, plus $\$35$). 
- If you bet $\$1$ on the first dozen (or second dozen, or third dozen) nonzero numbers and win, then your **net winnings** is $\$2$ (i.e. you get your original dollar back, plus another $\$2$. 

<img src="img/roulette_wheel.jpeg" width="330px">

It seems like there are so many ways to win!  In reality, some very careful probability theory was done by the game designers to ensure that there is not much difference in any particular payout.  We'll explore roulette both by simulation and by hand in this exercise. 

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3a (3 points)

We define **net winnings** to be the amount you win minus the amount you paid to play.

Let $X$ be a random variable that gives your **net winnings if you bet on red** and the roulette wheel is **spun once**.   

**i).**  What is the probability distribution (aka Probability Mass Function) of $X$?  Give your answer as a Markdown table.  

 - Tip:  To create Markdown tables that will compile when you create the PDF in this assignment, please refer to this syntax guide: https://www.markdownguide.org/extended-syntax/#tables.   Make sure you have a blank line before starting any tables, otherwise they won’t compile correctly in the PDF.



**ii).**  Calculate the **expected value** of your net winnings by betting on red.   

**iii).** Calculate the **variance** of your net winnings by betting on red.

Write up your full solution to all 3 parts in the SAME box below using Markdown and LaTeX (not code).  Show all steps fully justifying your answer.  
    


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3b (2 points)

Let's simulate this.  In the first code box below, write code to simulate one spin of a roulette wheel. Your code should take in a single variable `num_spins` and return an array of roulette spins of that length **WITHOUT USING A LOOP**.  Your output should be an array of strings in the form of the number then the color (i.e. `array(['18R', '00G',...])`) 

In the 2nd code box below, write code that takes the number of spins and either the color red or black as input, calculates winnings for each spin assuming you bet on that color for all spins, and then outputs the average winnings out of those spins.  

Then run the simulation for `num_spins` = 10,000,000  and compare to your answer from part A.

**To receive credit you must write your code such that all lines are visible in your PDF output.**

```python
def spin_roulette(num_spins): 
...
# Your code above this line

spin_roulette(5)

```

```python
def color_winnings(color='R', num_spins=100000):

...
# Your code above this line
    
print("E[Winnings] = {:.3f}".format(color_winnings(color="R", num_spins=int(1e7))))
```

```python
grader.check("q3b")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3c (3 points)

In Roulette you can bet on one of three "dozens" segments, called 1st 12, 2nd 12, and 3rd 12. They cover 1-12, 13-24, and 25-36, respectively.   If you bet $\$1$ on the first dozen (or second dozen, or third dozen) nonzero numbers and win, then your **net winnings is** $\$2$ (i.e. you get your original dollar back, plus another $\$2$. 


Let $Y$ be a random variable that gives your **net winnings** if you bet on any one of the three "dozen" nonzero numbers and the roulette wheel is spun once.  

i).  What is the probability distribution of $Y$?  Give your answer as a table. 

ii).  What is $E[Y]$?

iii). What is $Var[Y]$?

Write up your full solution to all 3 parts above in the SAME box below using LaTeX (not code).  Show all steps fully justifying your answer.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3d (2 points)

Write code to simulate `num_spins` spins, record the winnings for each spin if you bet on the first dozen nonzero numbers, and calculate the average winnings out of the total spins.

Then run the simulation for `num_spins` = 10,000,000  and compare to your answer from part C.  


```python
def dozen_winnings(num_spins):

...
# Your code above this line
    
print("E[Winnings] = {:.3f}".format(dozen_winnings(num_spins=int(1e7))))

```

```python
grader.check("q3d")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3e (3 points)
    
Recall, we showed in class that the expected value of your net winnings if you bet on any number is also $\$-\dfrac{1}{19}$.    

So  you're hopefully onto the pattern by now.  The net winnings in Roulette are designed so that the expected value of the net winnings is always $-\dfrac{1}{19}$ (i.e. on average the gambler loses money and the casino makes money).  

Since we define these in terms of your net winnings after betting $1, we can think of these as net winning odds.  

For example, since if you bet $\$1$ on the first dozen nonzero numbers and win, then your net winnings are $\$2$, we say the odds are 2 to 1 (denoted 2:1).   

The odds are 35:1 for landing on any particular number.  This means if you bet $\$1$, your net winnings will be $35.
 
 
Suppose the casino wanted to develop odds for a new bet in Roulette, where they allow you to bet on any set of 3 different numbers.   Let the odds for this new bet be $$x:1$$ ( This means if you bet $\$1$, your net winnings will be $\$x$).  

What should $\$x$ be so that the expected value of the net winnings is still $-\dfrac{1}{19}$?.

Show your work using LaTeX below.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3f (2 points)
    
Let's generalize this!   

Define a function $x(n)$ that describes the odds the casino should give for betting $\$1$ on any set of $n$ numbers if the casino wants to keep the expected value of the net winnings at $-\frac{1}{19}$ for $0 < n \leq 36$.   (For example, the odds for betting on any 3 different numbers should be set at $x(3)$ to $1$.  The odds for betting on any 4 different numbers should be set at $x(4)$ to $1$).   

Write your solution in the cell below:


_Type your answer here, replacing this text._

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p4'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4 (10 points)


Consider the discrete random variable $X$, with probability distribution given by $P(X=k) = 2ak^2-2ak$ for $k = 2, 3, 4$, where $a$ is a constant.


<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4a (8 points)

Answer all of the parts below in the SAME cell below using LaTeX.  Show all of your steps.  

**i**).  Determine the value of $a$ such that this defines a valid probability distribution.  Use that value for the rest of the problem.  

**ii**).  Calculate $P(X\leq3)$.

**iii**).  What is $E[X]$? (Show steps calculating this).

**iv**).  What is the standard deviation of $X$? (Show all steps calculating this).

Answer all of the parts above in SINGLE cell provided below using LaTeX.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answers to 4aiii and 4aiv, enter them below.  Note that this autograder test is not worth any points, it's just for you to check your work above to make sure you're on the right track. 

```python
expected_value = ...
standard_deviation= ...
```

```python
grader.check("q4_check")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4b (2 points)

Plot a histogram of the discrete probability distribution for $X$.  

Use the same plotting guidelines as shown in Problem 2: 
 - Set the bin widths to be equal to 1 
 - Add white lines between each bar 

**Be sure to include a title and label your x and y axes on your plot.**


```python
...
# Your code for the histogram above this line

```

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p5'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5 (7 points)

**Sampling With Replacement**


Consider a population in which a proportion $p$ of individuals are called "successes" (or 1, if you prefer) and the remaining proportion are rudely called "failures" (or 0).

We can sample from this population **at random with replacement**.  At random with replacement means we randomly choose one individual from the population and record whether or not they are a "success" or a "failure."  We then return that individual to the population before sampling again.  

(The reason why the individual is returned to the population before the next individual is drawn is to make sure the probability of selecting any particular individual remains the same in future draws).


If you draw a sample of size $n$ (where $n$ is some positive integer) at random with replacement from the population, then  probability of drawing $k$ successes and $n-k$ failures is $\binom{n}{k}p^k(1-p)^{n-k}$, for $0 \leq k \leq n$. To reduce writing, we will shorten "$k$ successes and $n-k$ failures" to "$k$ successes". In other words, "$k$ successes" means "exactly $k$ successes".

Since this random variable shows up in many applications it has an official mathematical name:  It is called a **binomial random variable** and its probability distribution is called a 
 **binomial distribution**.

To formalize here's how we write this in mathematical notation: 

 - We write $S\sim Bin(n, p)$ to represent that **the random variable $S$ is distributed as a Binomial Distribution with parameters $n$ and $p$.**  

 - The PMF (probability mass function) for $S$ is 
$P(S = k) = \binom{n}{k}p^k(1-p)^{n-k}$.

**Calculating binomial coefficients in Python**

Factorials and the *binomial coefficients* $\binom{n}{k} = \frac{n!}{k!(n-k)!}$ get large very quickly as $n$ gets large. One way to compute them is to use the `SciPy` module `special`. `SciPy` is a collection of Python-based software for math, probability, statistics, science, and engineering.

```python
from scipy import special
```

Below are some examples of `special.factorial`:

```python
special.factorial(5)

special.factorial(range(1, 6))
```

Traditionally, subsets of $k$ individuals out of a population of $n$ individuals are called *combinations*, and so `special.comb(n, k)` evaluates to $\binom{n}{k}$. 

Note that we will always use the term *subsets* to mean un-ordered sets. We will use *permutations* in situations where we need to keep track of the order in which the elements appear.


Look at the code and output below carefully (including types) to understand how these functions work.


https://docs.scipy.org/doc/scipy/reference/generated/scipy.special.comb.html



```python
special.comb(5, 3), special.factorial(5) / (special.factorial(3) * special.factorial(2))
```

```python
special.comb(5, range(6))
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5a (1 point)
Suppose you sample 100 times at random with replacement from a population in which 26% of the individuals are successes. Write a Python expression that evaluates to the chance that the sample has 20 successes. 

**Computational note:** Don't import any other libraries; just use the ones already imported and plug into the formula above. It's far from the best way numerically, but it is fine for the numbers involved in this Homework. 

```python
prob_5a = ...
prob_5a
```

```python
grader.check("q5a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5b (1 point)

Complete the cell with a Python expression that evaluates to an array whose elements are the chances of $k$ successes for $k = 0, 1, 2, \ldots, 100$. That is, `all_probs[k]` should contain $P(S = k)$, where $S$ is the random variable describing the number of successes drawn from the scenario outlined above.

```python
k = np.arange(101)
# Use 'k' above and np vectorization
all_probs = ...
np.round(all_probs, 4)
```

```python
grader.check("q5b")
```



Using the array of probability values we can plot a histogram of this discrete probability distribution as follows.

**NOTE:** We set the bin widths to be equal to 1 and add white lines between each value so that we can interpret area in the histogram as probabilities.

```python
fig, ax = plt.subplots()

ax.bar(k, all_probs, width=1, ec='white');
ax.set_axisbelow(True)
ax.grid(alpha=0.25)
plt.xlim(1,60)
#Always label your x-axis
plt.xlabel("k")
#Always label your y-axis
plt.ylabel("P(S=k)")

#Always include a title
plt.title("Binomial Distribution, n=100, p=0.26");


```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5c (1 point)


Complete the cell with an expression that evaluates to the chance that the number of successes in the sample is in the interval $26 \pm 10$ (inclusive on both sides). In other words, determine $P(16 \leq S \leq 36)$. **Hint: You might want to apply your Python indexing knowledge here.**

**Note: Please assign your answer to `prob_5c`. Because we use an autograder, please make sure your answer is a proportion between 0 and 1, not a percent between 0 and 100.**

```python
prob_5c = ...
prob_5c
```

```python
grader.check("q5c")
```

Consider (once again) a sample of size $n$ drawn at random with replacement from a population in which a proportion $p$ of the individuals are called successes.

Let $S$ be the random variable that denotes the number of successes in our sample. (As stated above, $S$ follows the **binomial distribution**.) Then, the probability that the number of successes in our sample is **at most** $s$ (where $0 \leq s \leq n$) is

$$P(S \leq s) = P(S = 0) + P(S = 1) + ... + P(S = s) = \sum_{k=0}^s \binom{n}{k}p^k(1-p)^{n-k}$$

We obtain this by summing the probability that the number of successes is exactly $k$, for each value of $k = 0, 1, 2, ..., s$.

We have a special name for the function $F(s) = P(S \leq s)$, for any given random variable $S$.   

It is called the **Cumulative Distribution Function (CDF)** of $S$.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5d (2 points)

Please fill in the function `cdf_S` which takes $n$, $p$, and $s$ and returns $P(S \le s)$ as defined above. 

If the inputs are invalid, for instance, if $p > 1$ OR $s > n$ then return 0. 



```python
def cdf_S(n, p, s):
    """ 
    returns the probability of S <= s
    Input n: sample size; p : proportion; s: number of successes at most"""
    ...
    # Your code above this line
```

```python
grader.check("q5d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5e (1 point)
    
Suppose in a particular an election, supporters of Candidate C are in a minority. Only 45% of the voters in the population favor the candidate.

Suppose a survey organization takes a sample of 200 voters at random with replacement from this population. Use `cdf_S` to write an expression that evaluates to the chance that a majority (more than half) of the sampled voters favor Candidate C.

```python
p_majority = ...
p_majority
```

```python
grader.check("q5e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5f (1 point)
Suppose each of five survey organizations takes a sample of voters at random with replacement from the population of voters in Part **e**, independently of the samples drawn by the other organizations. 

- Three of the organizations use a sample size of 200
- One organization uses a sample size of 300
- One organization uses a sample size of 400

Write an expression that evaluates to the chance that **in at least one of the five samples,** the majority of voters favor Candidate C. You can use any quantity or function defined earlier in this exercise.

```python
prob_1f = ...
prob_1f
```

```python
grader.check("q5f")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 5!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any additional websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 5 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW4 Autograded assignment in Gradescope.

* **Homework 5 Manually Graded**: Submit your hw05.PDF to the HW5 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Discord. 

```python
import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("hw05.ipynb")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw05.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for 2a (points: 1)
```
>>> assert isinstance(d_sup, list), 'Your answer must be a Python list, not a NumPy array or other type.'
```
```
>>> assert get_hash(len(d_sup)) == '1679091c5a880faf6fb5e6087eb1b2dc'
```
```
>>> assert get_hash(sorted(set(d_sup))) == '199ff5b613f5dc25dff99df513516bf9'
```


Tests for 2b (points: 1)
```
>>> for i in range(1, 7):
...     assert get_hash(round(d1_pmf[i], 4)) == 'f452927b0b2a755192117d956a8d47f6'
```
```
>>> assert get_hash(round(sum(d1_pmf), 3)) == 'e4c2e8edac362acab7123654b9e73432'
```


Tests for 2c (points: 1)
```
>>> assert isinstance(s_sup, list), 'Your answer must be a Python list, not a NumPy array or other type.'
```
```
>>> assert get_hash(len(s_sup)) == '6512bd43d9caa6e02c990b0a82652dca'
```
```
>>> assert get_hash(sorted(set(s_sup))) == '4511e38b335fb70fa03c080c37778523'
```


Tests for 2d (points: 1)
```
>>> assert get_hash(len(s_pmf)) == '6512bd43d9caa6e02c990b0a82652dca'
```
```
>>> assert get_hash(round(sum(s_pmf), 4)) == 'e4c2e8edac362acab7123654b9e73432'
```
```
>>> assert get_hash(s_pmf[4]) == 'e63ba2a42467cdb7252d28c6c05f2bba'
```
```
>>> assert get_hash(s_pmf[7]) == '97a5c78a55d017c4138b752f862e4fda'
```


Tests for q3b (points: 2)
```
>>> for spin in spin_roulette(10000):
...     assert spin[-1] in ('R', 'B', 'G')
...     assert int(spin[:-1]) >= 0 and int(spin[:-1]) <= 36
```
```
>>> assert get_hash(round(color_winnings(color='R', num_spins=int(10000000.0)), 2)) == 'e927fe5f5ad45bcd256f7039b3783ffb'
```
```
>>> import time
>>> start = time.time()
>>> res = spin_roulette(int(10000000.0))
>>> end = time.time()
>>> assert end - start < 1, 'Your solution is not fast enough! Try using list comprehension'
```


Tests for q3d (points: 2)
```
>>> assert get_hash(round(dozen_winnings(int(10000000.0)), 2)) == 'e927fe5f5ad45bcd256f7039b3783ffb'
```


Tests for q4_check (points: 0)
```
>>> assert get_hash(round(expected_value, 4)) == '804066a6a5042090a2a0b25c55f38e72'
```
```
>>> assert get_hash(round(standard_deviation, 5)) == '3877b6cb7d342cdbadd9d2fa54693b4a'
```


Tests for q5a (points: 1)
```
>>> assert prob_5a >= 0 and prob_5a <= 1
```
```
>>> assert get_hash(round(prob_5a, 5)) == '2a4e7dd819758fe6fd4db8fbc32f12e8'
```


Tests for q5b (points: 1)
```
>>> assert np.all(all_probs >= 0) and np.all(all_probs <= 1)
```
```
>>> k = np.arange(101)
>>> test_1b = special.comb(100, k) * 0.26 ** k * 0.74 ** (100 - k)
>>> assert np.all(np.isclose(test_1b, all_probs, atol=1e-05))
```


Tests for q5c (points: 1)
```
>>> assert prob_5c >= 0 and prob_5c <= 1
```
```
>>> assert get_hash(round(prob_5c, 4)) == '536a1fae1bafde3a6f3474dcd7e5e105'
```


Tests for q5d (points: 2)
```
>>> assert cdf_S(3, 0.4, 1) >= 0
```
```
>>> assert cdf_S(5, 0.6, 3) <= 1
```
```
>>> assert get_hash(round(cdf_S(2, 3, 4), 4)) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(round(cdf_S(10, 0.7, 4), 4)) == '2e0bfa827d51449a197fae6ab9f4804c'
```
```
>>> assert get_hash(round(cdf_S(5, 1, 1), 4)) == '30565a8911a6bb487e3745c0ea3c8224'
```
```
>>> assert get_hash(round(cdf_S(8, 0.7, 3), 4)) == 'e7cfe1112d37772ab690f04a56edee26'
```


Tests for q5e (points: 1)
```
>>> assert p_majority >= 0 and p_majority <= 1
```
```
>>> assert get_hash(round(p_majority, 4)) == '63cd3656f543ae3efc279fdd71e109d0'
```


Tests for q5f (points: 1)
```
>>> assert prob_1f >= 0 and prob_1f <= 1
```
```
>>> assert get_hash(round(prob_1f, 4)) == '22559430036a037ee145a651d1e91010'
```
