```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7:

## Due Date: Wednesday July 2nd by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Discord!


If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced

#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 5 | autograded & manual
Question 2 | 7 | manual
Question 3 | 6 | autograded
Question 4 | 7 | autograded & manual
Question 5 | 11 | autograded & manual
Question 6| 8 | autograded & manual
Question 7 | 6 | autograded
|Total | 50


## Recommended Readings for this HW:

- [CS109: Joint Distributions](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/joint/)
- [CS109: Marginalization](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/marginalization/)
- [CS109: Independent RV](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/independent_vars/)
- [CS109: Multinomial Distribution](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/multinomial/)
- [CS109: Covariance & Correlation](https://chrispiech.github.io/probabilityForComputerScientists/en/part3/correlation/)
- [CIT: 10.1 & 10.2:  Sampling and Empirical Distributions](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
- [LDS: Sec 2.2 - Target Population, Access and Sample](https://learningds.org/ch/02/data_scope_construct.html)
- [LDS: 2.5: Accuracy](https://learningds.org/ch/02/data_scope_accuracy.html)
- [LDS: 3.1 The Urn Model](https://learningds.org/ch/03/theory_urn.html)
- [LDS: 3.2 Example: Simulating Election Poll Bias and Variance](https://learningds.org/ch/03/theory_election.html)
- [Sampling](https://ds100.org/course-notes/sampling/sampling.html)

```python
import numpy as np
import pandas as pd
import math
import scipy.stats as stats
from scipy import special
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
from IPython.display import display, Latex, Markdown
import hashlib

plt.style.use('fivethirtyeight')
# Default plot configurations
plt.rcParams['figure.figsize'] = (16,8)
plt.rcParams['figure.dpi'] = 150
sns.set()

%matplotlib inline

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

<a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) |  [Problem 5](#p5) | [Problem 6](#p6) | [Problem 7](#p7) | 

               
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



# Part 1:  Joint Distributions, Covariance and Correlation

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1 (5 points)

Suppose $X$ and $Y$ are two discrete random variables with the following joint probability mass function:



$$
\begin{array}{|c|c|c|}
\hline
 & Y=1 & Y=2 \\
\hline
X=1 & \frac{1}{3} &  \\
\hline
X=2 & \frac{1}{6} & 0  \\
\hline
X=4 & \frac{1}{12} & \frac{1}{3}  \\
\hline
\end{array}
$$


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1a (1 point)
    
What is $P(X=1, Y=2)$?  

```python
prob_xequals1_and_yequals2 = ...
```

```python
grader.check("q1a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1b (1 point)
    
Find the marginal PMF for X, (i.e. find $P(X=x)$ for $x=1, 2, 4$).  

```python
prob_x_equals_1 = ...
prob_x_equals_2 = ...
prob_x_equals_4 = ...
```

```python
grader.check("q1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1c (1 point)
    
Find the marginal PMF for Y (i.e. find $P(Y=y)$ for $y = 1, 2$)

```python
prob_y_equals_1 = ...
prob_y_equals_2 = ...
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1d (2 points)
    
Are $X$ and $Y$ independent or dependent?  Fully justify your answer in the cell below using LaTeX and the mathematical definition of independence.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer to question 1d, enter a boolean value, True or False (note to receive credit on 1d you must mathematically justify your answer in the cell above). 

```python
x_y_independent = ...
```

```python
grader.check("q1d_check")
```

***

[Back to top](#top)

<a id='p2'></a>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2 (7 points)

The following parts are **NOT** related:


<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2a (2 points)
If Cov$(X,Y)=0$, what does this tell us about the random variables X and Y?  (Do not use the words "covariance" or "correlation" in your answer, instead explain in layman's terms how to interpret what this result means).  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2b (5 points)

Answer **all** of the parts below in the SAME cell below using Markdown and LaTeX.  Show all of your steps. Give your final answer as a fully simplified fraction

Given the following joint pmf for discrete random variables $X$ and $Y$:

$$
\begin{array}{|c|c|c|c|}
\hline
 & Y=0 & Y=1 & Y=2 \\
\hline
X=0 & \frac{1}{6} & \frac{1}{4} &\frac{1}{8} \\
\hline
X=1 & \frac{1}{8} & \frac{1}{6} &\frac{1}{6} \\
\hline
\end{array}
$$


**i).**  Calculate $Cov(X,Y)$.  
**ii).**  Calculate $\rho(X,Y)$   
 



_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 2bi, enter the answer you came up with (just the number) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
cov_xy = ...
```

```python
grader.check("q2bi_check")
```

To check your answer for Problem 2bii, enter the answer you came up with (just the number) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
p_xy = ...
```

```python
grader.check("q2bii_check")
```

<hr>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Simulating Multivariate Distributions

The outcome of the US presidential election in 2016 took many people and many pollsters by surprise. In this assignment we will carry out a simulation study / post mortem in an attempt to understand what happened.



***

[Back to top](#top)

<a id='p3'></a>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 3 (6 points)

#### Complete this [Multinomial Distribution; Sampling Video Assignment (click this link)](https://canvas.colorado.edu/courses/120673/assignments/2371677) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).

```python
from scipy import stats
```

```python
# Simulate a random sample (with replacement) of size 7, from a multinomial 
# distribution where the probability of selecting blue is 60%, green is 30% and red is 10%:

```

***

[Back to top](#top)

<a id='p4'></a>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4 (7 points)

In Pennsylvania, 6,165,478 people voted in the 2016 Presidential election.

Trump received 48.18% of the vote and Clinton recieved 47.46%.
This doesn't add up to 100% because other candidates received votes.
All together these other candidates received 100% - 48.18% - 47.46% = 4.36% of the vote.

The table below displays the proportion of Pennsylvania voters who voted for each candidate.




|   Voted for   |  Trump|    Clinton|    Other|
|-----------|-----------|-----------|---------|
| Probability      |   0.4818   | 0.4746  |   0.0436 |


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4a (1 point)
    
Suppose we pick a simple random sample of 20 of the 6,165,478 Pennsylvania (PA) voters. 

In the sample, let $N_T$ be the number of Trump voters, $N_C$ the number of Clinton voters, and $N_O$ the number of "other" voters. Then $N_T$, $N_C$, and $N_O$ are random: they depend on how the sample comes out. Such quantities are called "statistics".

Pick the correct option: $N_T + N_C + N_O$ is equal to

**(a)** 3

**(b)** 20

**(c)** 6,165,478

**(d)** a random quantity

*Note*: if your answer is (d), put  **ans_4a = 'd'**  for the purpose of grading. Passing the test below **does NOT** mean you answer is correct (there are hidden tests that will check this question when you submit your assignment). The test below only checks if your answer is in the correct format.

```python
ans_4a = ...
```

```python
grader.check("q4a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4b (1 point)
    
Pick the correct option.

A simple random sample (i.e. a sample without replacement) of 20 PA voters is like a sample drawn at random with replacement, because

**(a)** that's the definition of "simple random sample"

**(b)** there are only 3 categories of voters, which is small in comparison to 20

**(c)** there are only 20 people in the sample, which is less than $10\%$ of the total number of PA voters

**(d)** all PA voters are equally likely to be selected

*Note*: if your answer is (d), put **ans_4b = 'd'**  for the purpose of grading. Passing the test below **does NOT** mean you answer is correct (there are hidden tests that will check this question when you submit your assignment). The test below only checks if your answer is in the correct format.

```python
ans_4b = ...
```

```python
grader.check("q4b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4c (2 points)

Let $t$, $c$, and $o$ be any three non-negative integers. For the sample defined in Part **a**, the probability that the sample contains $t$ Trump voters, $c$ Clinton voters and $o$ "other" voters is denoted by $P(N_T = t, N_C = c, N_O = o)$.



Define a function `prob_sample_counts` that takes any three non-negative integers $t$, $c$, and $o$, and returns $P(N_T = t, N_C = c, N_O = o)$.  Note that the results of **Part b** imply that you can use results for sampling with replacement when calculating these probabilities.  
The probability is 0 for some choices of the arguments, and your function should return 0 in those cases.


**Hint: Use what you learned in the video lesson in Question 3.** 

**IMPORTANT NOTE:  To calculate any factorials use `special.factorial`    (do not use math.factorials as it will fail in gradescope)**


```python
def prob_sample_counts(t, c, o):
    """
    Input:
    t - number of votes for Trump
    c - number of votes for Clinton
    o - number of votes for Other
    Return: The probability of getting a such sample
    """
    ...
    # your code above this line
```

```python
grader.check("q4c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4d (1 point)

Check that your function determines a probability distribution by summing all the positive terms. You know what the sum should be. **Hint: You might need to iterate through certain values of c and t to compute the sum of all positive probabilities.**

```python
prob_4d = ...
prob_4d
```

```python
grader.check("q4d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4e (1 point)
    
Use `prob_sample_counts` to find the chance that the sample consists of 11 Trump voters, 8 Clinton voters, and 1 "other" voter.

```python
prob_4e = ...
prob_4e
```

```python
grader.check("q4e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 4f (1 point)
    
Write an expression that evaluates the chance that the sample contains at least 7 Trump voters, at least 7 Clinton voters, and at least 1 "other" voter. **Hint: You might find 4d useful.**

```python
prob_4f = ...
prob_4f
```

```python
grader.check("q4f")
```

[Back to top](#top)

<a id='p5'></a>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5 (11 points)


## Election Polling

Political polling is a type of public opinion polling that can at best represent a snapshot of public opinion at the particular moment in time. Voter opinion shifts from week to week, even day to day, as candidates battle it out on the campaign field.

Polls usually start with a "horse-race" question, where respondents are asked whom they would vote for in a head-to-head race if the election were tomorrow: Candidate A or Candidate B. The survey begins with this question so that the respondent is not influenced by any of the other questions asked in the survey. Some of these other questions are asked to help assess how likely is it that the respondent will vote. Other questions are asked about age, education, and sex in order to adjust the findings if one group appears overly represented in the sample.

Pollsters typically use [random digit dialing](https://en.wikipedia.org/wiki/Random_digit_dialing) to contact people.

Here are some important definitions:
- Population of Interest: All the people who will vote / voted in the 2016 US presidential election
- Sampling Frame: Anyone who has a phone (they need to be reached by random digit dialing). In addition, survey respondents are excluded if they are deemed unlikely to vote or not elligible for voting in the upcoming election.

### How might the sampling frame differ from the population?

After the fact, many experts have studied the 2016 election results. For example, according to the American Association for Public Opinion Research (AAPOR), predictions made before the election were flawed for three key reasons:

1. voters changed their preferences a few days before the election
2. those sampled were not representative of the voting population, e.g., some said that there was an overrepresentation of college graduates in some poll samples 
3. voters kept their support for Trump to themselves (hidden from the pollsters)



The 1st and 3rd biases (voters changing preference and voters hiding their preference) can only be determined after the election. It might be possible to assess this with a post election survey where people were asked about these things.   

In the rest of this assignment we will focus on analyzing the 2nd  type of bias:


+ Question 5: We will carry out a study of the sampling error when there is no bias. In other words, we will try to compute the chance that random polls get the election result wrong even if we collect our sample in a manner that is completely correct. In this case, any **failure of our prediction is due entirely to random chance**.
+ Question 6: We will carry out a study of the sampling error when there is bias of the second type from the list above. In other words, we will try to compute the chance that random polls get the election result wrong if we have a small systematic bias. In this case, any failure of our prediction is due to a **combination of random chance and our bias**.



### How large was the sampling error?

In some states the race was very close, and it may have been simply sampling error, i.e., random chance that the majority of the voters chosen for the sample voted for Clinton.

A 2- or 3-point polling error in Trump’s favor (typical error historically) would likely be enough to tip the Electoral College to him.

One year after the 2016 election, Nate Silver (https://substack.com/@natesilver) wrote in
*The Media Has A Probability Problem* that
the "media’s demand for certainty -- and its lack of statistical rigor -- is a bad match for our complex world."
FiveThirtyEight forecasted that Clinton had about a 70 percent chance of winning.  

We will first carry out a simulation study to assess the impact of the sampling error on the predictions.


## The Electoral College

The US president is chosen by the Electoral College, not by the
popular vote. Each state is alotted a certain number of 
electoral college votes, as a function of their population.
Whomever wins in the state gets all of the electoral college votes for that state.

There are 538 electoral college votes (hence the name of the website Nate Silver started, FiveThirtyEight: https://projects.fivethirtyeight.com/polls/).

Pollsters correctly predicted the election outcome in 46 of the 50 states. 
For these 46 states Trump received 231 and Clinton received 232 electoral college votes.

The remaining 4 states accounted for a total of 75 votes, and 
whichever candidate received the majority of the electoral college votes in these states would win the election. 

These states were Florida, Michigan, Pennsylvania, and Wisconsin.

|State |Electoral College Votes|
| --- | --- |
|florida | 29 |
|michigan | 16 |
|pennsylvania | 20 |
|wisconsin | 10|

For Donald Trump to win the election, he had to win either:
* Florida + one (or more) other states
* Michigan, Pennsylvania, and Wisconsin


In the final results the electoral margins were very narrow in these four states, as seen below:


|State | % Trump |   % Clinton | Total Voters |
| --- | --- |  --- |  --- |
|florida | 49.02 | 47.82 | 9,419,886  | 
|michigan | 47.50 | 47.27  |  4,799,284|
|pennsylvania | 48.18 | 47.46 |  6,165,478|
|wisconsin | 47.22 | 46.45  |  2,976,150|

Those narrow electoral margins can make it hard to predict the outcome **given the sample sizes** that the polls used. 

***


    
### Simulation Study of the Sampling Error

Now that we know how people actually voted, we can carry
out a simulation study that imitates the polling from 2016.

**NOTE**: We know the results of this election, as it took place in 2016. Because we know Trump won:
- Anytime Clinton is predicted to win in the simulated poll, that prediction is incorrect
- Anytime Trump is predicted to win in the simulated poll, that prediction is correct

Our ultimate goal in this problem is to use the actual voting percentages to simulate polling results and **understand the chance that our poll we will incorrectly call the election for Hillary Clinton even if our sample was collected with absolutely no bias**.



<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5a (2 points)

For your convenience, the actual results of the vote in the four pivotal states is repeated below:

|State |% Trump |   % Clinton | Total Voters |
| --- | --- |  --- |  --- |
|florida | 49.02 | 47.82 | 9,419,886  | 
|michigan | 47.50 | 47.27  |  4,799,284|
|pennsylvania | 48.18 | 47.46 |  6,165,478|
|wisconsin | 47.22 | 46.45  |  2,976,150|


Using the table above, write a function `draw_state_sample(N, state)` that returns a sample with replacement of N voters from the given state, using the percentages given in the table above. Your result should be returned as a list, where the first element is the number of Trump votes, the second element is the number of Clinton votes, and the third is the number of Other votes. For example, `draw_state_sample(1500, "florida")` could return `[727, 692, 81]`. You may assume that the state name is given in all lower case.  

**Note**: Use the given rounded probability. DO NOT use an exact fraction or a re-rounded probability as this will make the autograder fail. 

**Hint:** Use  `scipy.rvs()` to randomly sample from an appropriate discrete joint distribution.  


```python
def draw_state_sample(N, state):
    ...
```

```python
draw_state_sample(1500, "florida")
```

```python
grader.check("q5a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5b (1 point)

Now, create a function `trump_advantage` that takes in list of sample of votes (like the one returned by `draw_state_sample`) and returns the difference in the proportion of votes between Trump and Clinton. For example `trump_advantage([100, 60, 40])` should return `0.2`, since Trump had 50% of the votes in this sample and Clinton had 30%.  In contrast, `trump_advantage([60, 100, 40])` should return `-0.2`

```python
def trump_advantage(voter_sample):
    ...
```

```python
grader.check("q5b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5c (2 points)

Simulate Trump's advantage across 100,000 random samples of 1500 voters for the state of Pennsylvania and store the results of each simulation in an `np.array` called `simulations`. 

That is, `simulations[i]` should be Trump's proportion advantage for the `i+1`th random sample of 1500 voters.

```python
simulations = ...

```

```python
grader.check("q5c")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5d (3 points)

i).  Make a **frequency** histogram of `simulations`.   This is a histogram of the sampling distribution of Trump's proportion advantage in Pennsylvania. 

Hint: You should use the [`plt.hist`](https://matplotlib.org/3.1.1/api/_as_gen/matplotlib.pyplot.hist.html) function in your code.

Make sure to include a title as well as axis labels. You can do this using `plt.title`, `plt.xlabel`, and `plt.ylabel`.

ii).  Based on your simulation, what is the probability that a random sample of 1500 will correctly predict that Trump wins Pennsylvania? (i.e. what proportion of these simulations predict a Trump victory?)  Assign your answer to `prob_penn_1500_random_correct`


```python
# Part (i):
...
# your code for the histogram above here.  The code below plots a red marker at the mean:
plt.scatter(simulations.mean(), -1, marker='^', color='red', s=500)
```

```python
# Part (ii):
prob_penn_1500_random_correct = ...

prob_penn_1500_random_correct
```

<!-- END QUESTION -->


**NOTE**: We know the results of this election, as it took place in 2016. Because we know Trump won:
- Anytime Clinton is predicted to win in the simulated poll, that prediction is incorrect
- Anytime Trump is predicted to win in the simulated poll, that prediction is correct


Thus your result above allows us to calculate the chance that we will incorrectly call the election for Hillary Clinton in Pennsylvania based on the results of this random sample of 1500 even if our sample was collected with absolutely no bias:



```python
prob_penn_1500_random_incorrect = 1-prob_penn_1500_random_correct
prob_penn_1500_random_incorrect
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5e (2 points)

Now write a function `trump_wins(N)` that creates a random sample of N voters for **each of the four crucial states** (that is N voters from Florida, N voters from Michigan, N voters from Pennsylvania, and N voters from Wisconsin) and returns 1 if Trump is predicted to win based on these samples and 0 if Trump is predicted to lose.   
(Hint: use `draw_state_sample` from above) 

Recall that for Trump to win the election, he must either:
* Win the state of Florida and 1 or more other states
* Win Michigan, Pennsylvania, and Wisconsin

Then repeat 100,000 simulations of the election, i.e. call `trump_wins(1500)` 100,000 times and calculate what proportion of these simulations predict a Trump victory. Give your answer as a float stored in `proportion_trump`.

This number represents the percent chance that samples of 1500 from each of the 4 states will correctly predict Trump's victory *if the samples were collected with absolutely no bias*. 


```python
def trump_wins(N):
    ...
    
proportion_trump = ...
proportion_trump
```

```python
grader.check("q5e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 5f (1 point)
    
Based on your work above, what is the chance that we will incorrectly call the election for Hillary Clinton based on the results of random samples of 1500 people in these 4 battleground states **even if our samples were collected with absolutely no bias?**    

```python
prob_election_1500_random_incorrect = ...
prob_election_1500_random_incorrect
```

```python
grader.check("q5f")
```

***

[Back to top](#top)

<a id='p6'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6 (8 points)
    
### Simulation Study of Selection Bias

According to [an article](https://blogs.lse.ac.uk/usappblog/2018/02/01/better-poll-sampling-would-have-cast-more-doubt-on-the-potential-for-hillary-clinton-to-win-the-2016-election/#Author) by Grotenhuis, Subramanian, Nieuwenhuis, Pelzer and Eisinga:

> In a perfect world, polls sample from the population of voters, who would state their political preference perfectly clearly and then vote accordingly.

That's the simulation study that we just performed. 


It's difficult to control for every source of selection bias.
And, it's not possible to control for some of the other sources of bias.

Next we investigate the effect of small sampling bias on the polling results in these four battleground states.  

Throughout this problem, we'll examine the impacts of a 0.5 percent bias in favor of Clinton in each state. Such a bias has been suggested because highly educated voters tend to be more willing to participate in polls.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6a (2 points)

Throughout this problem, adjust the selection of voters so that there is a 0.5% bias in favor of Clinton in each of these states. 

For example, in Pennsylvania, Clinton received 47.46\% of the votes and Trump 48.18\%. Increase the population of Clinton voters to 47.46\% + 0.5\%  and correspondingly decrease the percent of Trump voters. 


Simulate Trump's advantage across 100,000 simple random samples of 1500 voters for the **state of Pennsylvania** and store the results of each simulation in an `np.array` called `biased_simulations`.

That is, `biased_simulation[i]` should hold the result of the `i+1`th simulation.

That is, your answer to this problem should be just like your answer from **Question 5c**, but now using samples that are biased as described above.


```python
def draw_biased_state_sample(N, state):
    
    ...
    
biased_simulations = ...
```

```python
grader.check("q6a")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6b (1 point)

Create a plot of **overlaid DENSITY** histograms of the following:
 - The new sampling distribution of Trump's proportion advantage in Pennsylvania using these biased samples 
  - The sampling distribution of the unbiased samples from **Question 5d** (plotted as a density, not a frequency histogram)
  
Include 2 markers (of different colors) with the sample means for each distribution (see **5d** for code how to do this).  The colors of the markers should correspond to the colors of the density histograms.

Make sure to give your plot a title, label the x and y axes and include a legend.
Use the parameter `alpha` to adjust the transparency of each histogram.




```python
...
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6c (1 point)

Use your result above to calculate the chance that we will incorrectly call the election for Hillary Clinton in Pennsylvania based on the results of this random sample of 1500 if our sample was collected with this $0.5\%$ bias:


```python
prob_penn_1500_biased_incorrect = ...
prob_penn_1500_biased_incorrect
```

```python
grader.check("q6c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6d (2 points)

Now perform 100,000 simulations of $1500$ voters from **each of the four states** (1500 voters from Pennsylvania, 1500 voters from Michigan, etc)  and return the proportion of these simulations that result in a Trump victory. This is the same fraction that you computed in **Question 5e**, but now using your biased samples.

Give your answer as `proportion_trump_biased`.


```python
def trump_wins_biased(N): 
    ...

proportion_trump_biased = ...
proportion_trump_biased
```

```python
grader.check("q6d")
```

The output from the above function represents the chance that a sample biased 0.5% in Hillary Clinton's favor will correctly predict Trump's victory. 

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 6e (2 points)

Summarize the findings from these simulations.


**Recall**: We know the results of this election, as it took place in 2016. Because we know Trump won:
- Anytime Clinton is predicted to win in the simulated poll, that prediction is incorrect
- Anytime Trump is predicted to win in the simulated poll, that prediction is correct


i).  Based on your simulations, what was the **probability** that polls using **unbiased** samples of $1500$ people from each state would incorrectly predict Clinton would win in 2016? (This is called **chance error**).  Many people, even well educated ones, assume that this number should be $0\%$. After all, how could a non-biased sample be wrong?   Give a mathematical explanation as to why it isn't $0\%$ (or close to $0\%$).   This is the type of incredibly important intuition we hope to develop in you throughout this class and your future data science coursework.


ii).  Based on your simulations, what was the **probability** that polls using **biased** samples of $1500$ people from each state would incorrectly predict Clinton would win in 2016?  How different is this from your answer in part (i)?   Recall, we only biased the samples by $0.5\%$. However, even a bias this small in the percentages can lead to a much larger probability of error in prediction of the final result.

_Type your answer here, replacing this text._

<!-- END QUESTION -->


***

[Back to top](#top)

<a id='p7'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 7 (6 points)

Would increasing the sample size have helped?


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 7a (1 point)


Let's try increasing the sample size.  Try a sample size of 5,000 (in each of the 4 states) and run 100,000 simulations of a sample with replacement. What proportion of the 100,000 times is Trump incorrectly predicted to lose the election:
 - In the unbiased setting?
 - In the biased setting?

Give your answers as `high_sample_size_unbiased_proportion_trump` and `high_sample_size_biased_proportion_trump`.



*Note: Don't worry if your solution takes a while to run.* 

```python
high_sample_size_unbiased_proportion_incorrect = ...
high_sample_size_biased_proportion_incorrect = ...

print(high_sample_size_unbiased_proportion_incorrect, high_sample_size_biased_proportion_incorrect)

```

```python
grader.check("q7a")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 7b (2 points)

Compare your observations from **7a** to your observations **6e**.  

i).  After increasing the sample sizes, did the probability of incorrectly predicting the election winner (i.e. sampling error) increase or decrease for the biased versus unbiased samples?  

ii). What do these changes imply about the impact of sample size on the sampling error in unbiased vs biased samples? 

Write your answer to both parts in the cell below: 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 7c (2 points)

Is it possible to correctly predict Trump's victory with less than $1\%$ error using **unbiased sampling?**  Rerun the simulation (in each of the 4 states) with increasing sample sizes and 100,000 simulations to determine if you can find an approximate minimum sample size (it doesn't have to be exact) such that the probability of correctly predicting Trump's victory is at least $99\%$ (assuming your sample is unbiased).  



```python
...
# your code above this line.
# output the number of samples you used to get to at least 99% accuracy.
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 7d (1 point)

Is it possible to correctly predict Trump's victory with less than $1\%$ error using **biased sampling?**  Use the code cell below to rerun the simulation (in each of the 4 states) with increasing sample sizes.  What happens to the probability of correctly predicting Trump's victory as you increase the sample size?  Explain in the markdown cell below.



_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
print("10000 runs gives", np.mean([trump_wins_biased(10000) for i in range(100000)]))
print("20000 runs gives",np.mean([trump_wins_biased(20000) for i in range(100000)]))
print("30000 runs gives",np.mean([trump_wins_biased(30000) for i in range(100000)]))
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 7!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any additional websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 7 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW7 Autograded assignment in Gradescope.

* **Homework 7 Manually Graded**: Submit your hw07.PDF to the HW6 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python
import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("hw07.ipynb")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw07.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1a results:
        q1a - 1 result:
            ❌ Test case failed
            Trying:
                assert get_hash(round(prob_xequals1_and_yequals2, 4)) == 'f71f055698f44ad6b4d99ff0d9adfb25'
            Expecting nothing
            **********************************************************************
            Line 1, in q1a 0
            Failed example:
                assert get_hash(round(prob_xequals1_and_yequals2, 4)) == 'f71f055698f44ad6b4d99ff0d9adfb25'
            Exception raised:
                Traceback (most recent call last):
                  File "/opt/conda/lib/python3.10/doctest.py", line 1350, in __run
                    exec(compile(example.source, filename, "single",
                  File "<doctest q1a 0[0]>", line 1, in <module>
                    assert get_hash(round(prob_xequals1_and_yequals2, 4)) == 'f71f0
```

Tests for q1a (points: 1)
```
>>> assert get_hash(round(prob_xequals1_and_yequals2, 4)) == 'f71f055698f44ad6b4d99ff0d9adfb25'
```


Tests for q1b (points: 1)
```
>>> assert get_hash(round(prob_x_equals_1, 4)) == 'bfae8f3038fb5eb9c961af6e8e8ab859'
```
```
>>> assert get_hash(round(prob_x_equals_2, 4)) == 'f452927b0b2a755192117d956a8d47f6'
```
```
>>> assert get_hash(round(prob_x_equals_4, 4)) == 'bfae8f3038fb5eb9c961af6e8e8ab859'
```


Tests for q1c (points: 1)
```
>>> assert get_hash(round(prob_y_equals_1, 4)) == '9fe424617e894efd62d2a628dd86e776'
```
```
>>> assert get_hash(round(prob_y_equals_2, 4)) == 'bfae8f3038fb5eb9c961af6e8e8ab859'
```


Tests for q1d_check (points: 0)
```
>>> assert get_hash(x_y_independent) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q2bi_check (points: 0)
```
>>> assert get_hash(round(cov_xy, 4)) == '7402532294a1062840be7cfcaf7496dd'
```


Tests for q2bii_check (points: 0)
```
>>> assert get_hash(round(p_xy, 4)) == '1bc01bc3bdacbbc89510044889faad30'
```


Tests for q4a (points: [0])
```
>>> assert ans_4a in ['a', 'b', 'c', 'd']
```


Tests for q4b (points: [0])
```
>>> assert ans_4b in ['a', 'b', 'c', 'd']
```


Tests for q4c (points: 2)
```
>>> assert 0 <= prob_sample_counts(2, 5, 13) <= 1
```
```
>>> assert 0 <= prob_sample_counts(5, 1, 14) <= 1
```
```
>>> assert 0 <= prob_sample_counts(10, 8, 2) <= 1
```
```
>>> assert 0 <= prob_sample_counts(6, 12, 2) <= 1
```
```
>>> assert get_hash(prob_sample_counts(6, 12, 3)) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(prob_sample_counts(20, 1, 3)) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(prob_sample_counts(6, 6, 2)) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(round(prob_sample_counts(6, 12, 2), 5)) == '95395938a3043704f9a63ddfeef0e1a0'
```
```
>>> assert get_hash(round(prob_sample_counts(13, 7, 0), 5)) == '274cb79c3c312edfbd68b64f164a2721'
```
```
>>> assert get_hash(round(prob_sample_counts(13, 5, 2), 5)) == 'febd7766501e80909143f02ed9180628'
```


Tests for q4d (points: 1)
```
>>> assert get_hash(round(float(prob_4d), 4)) == 'e4c2e8edac362acab7123654b9e73432'
```


Tests for q4e (points: 1)
```
>>> assert get_hash(round(prob_4e, 4)) == '719bef9aa5ee3b155865212824803fd2'
```


Tests for q4f (points: 1)
```
>>> assert get_hash(round(prob_4f, 4)) == 'a4a33b8dc1b9e8f39da2e6aff66c24f0'
```


Tests for q5a (points: 2)
```
>>> assert get_hash(len(draw_state_sample(1500, 'florida'))) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
```
>>> assert get_hash(np.sum(draw_state_sample(1500, 'michigan'))) == 'cfa5301358b9fcbe7aa45b1ceea088c6'
```
```
>>> q5a_penn = draw_state_sample(1500, 'pennsylvania')
>>> trump_win_penn = (q5a_penn[0] - q5a_penn[1]) / 1500
>>> assert abs(trump_win_penn - 0.007) <= 0.12
```
```
>>> N = 1000
>>> assert all([0 <= x / N and x / N <= 1 for x in draw_state_sample(N, 'wisconsin')])
```


Tests for q5b (points: [0, 0.5, 0.5])
```
>>> assert -1 < trump_advantage(draw_state_sample(1500, 'wisconsin')) < 1
```
```
>>> assert np.isclose(trump_advantage([100, 60, 40]), 0.2)
```
```
>>> assert get_hash(round(trump_advantage([10, 30, 10]), 4)) == 'd1ab6103a7baf2fcf7a1611ac2299261'
```


Tests for q5c (points: [0, 1, 1])
```
>>> assert len(simulations) == 100000
```
```
>>> assert get_hash(np.sum([-1 < x < 1 for x in simulations])) == '14ee22eaba297944c96afdbe5b16c65b'
```
```
>>> assert abs(np.mean(simulations) - 0.007) <= 0.016
```


Tests for q5e (points: [0, 2])
```
>>> assert trump_wins(1000) in [0, 1]
```
```
>>> assert get_hash(abs(proportion_trump - 0.695) <= 0.02) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q5f (points: 1)
```
>>> assert get_hash(abs(1 - proportion_trump - 0.3) <= 0.03) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q6a (points: 2)
```
>>> assert get_hash(len(draw_biased_state_sample(1000, 'wisconsin'))) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
```
>>> assert get_hash(np.sum(draw_biased_state_sample(1000, 'michigan'))) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
```
```
>>> q6_1 = draw_biased_state_sample(3000, 'florida')
>>> assert q6_1[0] > q6_1[2] and q6_1[1] > q6_1[2]
```
```
>>> assert len(biased_simulations) == 100000
```
```
>>> assert get_hash(np.sum([-1 < x < 1 for x in biased_simulations])) == '14ee22eaba297944c96afdbe5b16c65b'
```
```
>>> assert abs(np.mean(biased_simulations) + 0.003) <= 0.12
```


Tests for q6c (points: 1)
```
>>> assert abs(prob_penn_1500_biased_incorrect - 0.55) <= 0.025
```


Tests for q6d (points: 2)
```
>>> assert abs(proportion_trump_biased - 0.466) <= 0.04
```
```
>>> assert get_hash(sorted(set([trump_wins_biased(1500) for i in range(100)]))) == '264884439b70ab09a86bc848421c6de6'
```


Tests for q7a (points: 1)
```
>>> assert np.abs(high_sample_size_unbiased_proportion_incorrect - 0.171) <= 0.02
```
```
>>> assert np.abs(high_sample_size_biased_proportion_incorrect - 0.555) <= 0.04
```
