```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08.ipynb")
```

# Homework 8:  Central Limit Theorem & Hypothesis Tests



## Due Date: Friday July 4th by 11:59 PM MT on Gradescope


## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Entering a HW question as a prompt into Generative AI and then either directly using or modifying the solution.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Discord!


#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown

Question | Points | Grading Type
--- | --- | ---
Question 1 | 9 | manual
Question 2 | 3 | autograded
Question 3a-g  | 7 | autograded
Question 3h | 2 | manual
Question 3i-j | 2 | autograded
Question 4a | 1 | manual
Question 4b | 1 | autograded
Question 4c | 2 | manual
Question 4d | 1 | autograded
Question 4e | 1 | manual
Question 5 | 5 | manual
Question 6a-g | 9 | autograded
Question 6h | 4 | manual
Question 6i | 1 | autograded
Question 6j | 2 | manual
|Total|50


## HW 8 Learning Objectives:

 - Explain the difference between parameters and statistics 

 - Explain the difference between a population distribution, a sample distribution and the sampling distribution of a statistic

 - Simulate the empirical sampling distribution of a statistic

 - Apply and interpret the Central Limit Theorem for sums and means (what is it used for?  What conditions must hold to use it?)

 - Explain when to use hypothesis tests

 - Implement hypothesis tests comparing a sample to a model and interpret results

 - Implement a hypothesis test to determine if a sample is consistent with a specific multinomial distribution

 - Define null and alternative hypotheses and explain guidelines for how to define them.

 - Define the test statistic and explain what it is used for.  Provide guidelines for types of test statistics based on your null and alternative hypotheses.

 - Define the significance level and explain what it is used for.

 - Define the p-value and interpret it both mathematically as well as visually.

 - Describe the difference between an empirical p-value and a theoretical p-value

 - Describe the types of conclusions you can make in hypothesis testing

 - Implement A/B hypothesis tests and interpret conclusions

 - Explain how and when to use permutations to simulate under the null

 - Explain when you can make causal conclusions from an A/B test

**Assigned Readings On These Topics from Class Schedule**: 



* [The Central Limit Theorem](https://inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html?highlight=central+limit+theorem)
* [Variability of the Sample Mean](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html)
  
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html) Chapter 11 - All

* [Comparing Two Samples](https://inferentialthinking.com/chapters/12/Comparing_Two_Samples.html): Chapter 12 - All


## <a id='top'></a>
---
**Shortcuts:**  [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) | [Problem 5](#p5) | [Problem 6](#p6) |
---
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

```python
import hashlib
import pandas as pd
import scipy.stats as stats
# These lines import the Numpy and Datascience modules.
import numpy as np
from matplotlib.ticker import PercentFormatter

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('seaborn-v0_8')
import warnings
warnings.simplefilter('ignore', FutureWarning)

from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 1  ##

The file `data/income_data.csv` contains Age and Income information from a fictitious city in the Midwest with 5000 residents.   Consider this data to be our population data.

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1a (2 pts) ## 

Load the data into a pandas DataFrame called `dfIncome`, calculate the population mean of Income and then  make a **density histogram** of the Distribution of the Income data **with 15 bins**.   Include a title for your plot and label the x-axis (we have provided a label for the y-axis).
Note we have included code to mark where the population mean lies on the histogram.

```python
dfIncome = ...

mean_income = ...

print("Population income mean is", mean_income)

...
# write code to plot density histogram above this line

plt.ylabel("Probability Per Dollar") # Since this is a density histogram, the y-units are probability per unit x
#Add a triangle marker to indicate where the population mean is
plt.scatter(mean_income, -.0000001, marker='^', color='red', s=300)

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1b (1 pts) ## 

Describe the shape of the Income distribution (i.e. comment on both modality and skew)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1ci (2 pts) ## 

Write a function to collect a random sample of size `sample_size` with replacement from `dfIncome`  and plot the **density histogram** of the empirical distribution of the income for your sample.  Use 15 bins for your histogram and set the x-axis range to be from (0,210000).   (Hint: use the dataframe method `.sample()`). Include a title and label for both axes.          
Then run the cells provided below to output 3 separate distributions for sample sizes of $10$, $100$ and $1000$. 


```python
def income_sample(df, sample_size):
    
...
# rest of code above this line
```

```python
income_sample(dfIncome,10)
```

```python
income_sample(dfIncome,100)
```

```python
income_sample(dfIncome,1000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1cii (1 pt) ## 

What happens to the shape of the empirical sample distributions of income as you increase the sample size, and how does it relate to the population distribution? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1d (2 pts) ## 

If we want to estimate the **mean** of the population we can draw a sample from the population and compute the sample mean.  As we learned in class, since samples can vary, the sample mean can vary and thus it is a random variable and has its own distribution.


i).  Complete the function `income_sample_mean` below to randomly sample `sample_size` rows from `dfIncome` with replacement and return the sample mean of income for that sample. 


ii).  Complete the function `income_sample_dist` below to simulate `num_simulations` of randomly sampling `sample_size` rows from `dfIncome` with replacement and calculate the sample mean of income for each sample.
Store the sample means in an np.array called `means`.  The function should output a **density** histogram of the empirical sample mean income distribution.    On the histogram, include two **triangle** markers on the histogram like in part 1a: A red one for the population mean (that you calculated in part 1A) and a yellow one for the mean of the `num_simulations` sample mean estimates.
Use 25 bins, do not set a range, and keep the x-axis limits fixed across all plots for comparison.
Include a title and labels for the x and y axes.

Then run the cells provided below to output 3 separate distributions for `num_simulations`=1000 and `sample_size` = $10$, $100$ and $1000$



```python
def income_sample_mean(df, sample_size):
    ...
```

```python
def income_sample_dist(df, sample_size, num_simulations):
    means = ...
      # 'means' stores "num_simulations" means from samples of size "sample_size"
    plt.xlim([20000,120000])
    ...
    
# Your code for part (ii) above
```

```python
income_sample_dist(dfIncome, 10, 1000)
```

```python
income_sample_dist(dfIncome, 100, 1000)
```

```python
income_sample_dist(dfIncome, 1000, 1000)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 1e (1 pt) ## 

Describe the shapes of the empirical sample mean distributions (comment on their modality and skew compared to the modality and skew of the population distribution).
What happens to the mean and standard deviations of these distributions as you increase the sample size?
What is the name of the theorem that explains what you are observing?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p2'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2 

A hardware store receives a shipment of 10,000 bolts that are supposed to be 12 cm long. 
The mean of this shipment of 10,000 bolts is indeed 12 cm, and the standard deviation is 0.2 cm.   

For the following questions, determine if you have enough information to answer.  If you do, then use built-in python functions to calculate the probability. If not, enter -1 for your answer.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2a (1 pt)
    
What is the probability that a randomly chosen bolt is less than 10 cm long?

```python
q2a_answer = ...
```

```python
grader.check("q2a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>  Question 2b (2 pts)

For quality control, the hardware store chooses 100 bolts at random to measure.  They will declare the shipment defective and return it to the manufacturer if the average length of 100 bolts is less than 11.97 cm or greater than 12.04 cm.    Find the probability that the shipment is found satisfactory (i.e. not defective).

```python
q2b_answer = ...
```

```python
grader.check("q2b")
```

***

[Back to top](#top)

<a id='p3'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3 ## 

The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student from Loveland, Colorado who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would be on the other side and would flip a fair coin. Depending on how the coin landed, she would put out either her left hand or her right hand. The TT practitioner would then have to answer which hand Emily put out. If a practitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, the practitioner picked the correct hand 44% of the time. 

Emily's main goal here was to test whether or not the TT practitioners' guesses were random, like the flip of a coin. In most medical experiments, this is the norm. We want to test whether or not the treatment has **any** effect, *not* whether or not the treatment actually worked.
We will now begin to formulate this experiment in terms of the terminology we learned in this course. 

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3a (1 pt) ## 

Which of the following is Emily's null hypothesis? 

 - (a):  Practitioners are answering independently and better than if they were guessing at random (i.e. the probability of any practitioner correctly answering is greater than 50%)
 
 - (b): Practitioners are answering independently and randomly (i.e. the probability of any practitioner correctly answering is 50%).  
  
 - (c):  Practitioners are answering independently and worse than if they were guessing at random (i.e. the probability of any practitioner correctly answering is less than 50%)
 
Assign `null_hypothesis` to the correct letter (enter it as a lower case string, for example 'a').  




*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.


```python
null_hypothesis = ...
```

```python
grader.check("q3a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3b (1 pt) ## 

Given that we want to test whether or not the treatment has **any** effect, *not* whether or not the treatment actually worked, what is the alternative hypothesis?

Assign `alt_hypothesis` to the correct letter (enter it as a lower case string)

 - (a):  Practitioners are answering independently and better than if they were guessing at random (i.e. the probability of any practitioner correctly answering is greater than 50%)
 
  
 - (b):  Practitioners are answering independently and worse than if they were guessing at random (i.e. the probability of any practitioner correctly answering is less than 50%)
 
 - (c): Practitioners are answering independently and not randomly (i.e. the probability of any practitioner correctly answering is not 50%).  






```python
alt_hypothesis = ...
```

```python
grader.check("q3b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3c (1 pt) ## 

Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's **null hypothesis**, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is a number between 0 and 1. 


```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q3c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3d (1 pt) ## 

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 


We usually use a statistic to help determine which model the evidence points towards. One of the test statistics below is appropriate for testing the null and alternative hypotheses you indicated above. Assign the variable `valid_test_stat` to the letter  (lower case  string) corresponding to the correct test statistic for this scenario. 



 - (a): 50 minus the actual percent correct
 - (b): the actual percent correct minus 50 
 - (c): 44 minus the actual percent correct
 - (d): the actual percent correct minus 44
 - (e): the sum of 50 and the actual percent correct
 - (f): the absolute value of the difference between 50 and the actual percent correct
 - (g): the absolute value of the difference between 44 and the actual percent correct


```python
valid_test_stat = ...
valid_test_stat
```

```python
grader.check("q3d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3e (1 pt) ## 

Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Question 3d. Assume that the argument takes in proportions, but  **return your answer as a percentage**   (i.e. an answer of 0.70 should be output as $70$)

*Hint:* Remember we are asking for a **percentage**, not a proportion. So int() will not work! 


```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q3e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3f (1 pt) ## 

Use your newly defined function from 3e to calculate the **observed test statistic** from Emily's experiment. 

```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q3f")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3g (1 pt) ## 

Next we want to explore if this observed statistic is consistent with what we expect to see under Emily’s null hypothesis.  To do this, we will need to simulate test statistics under the null hypothesis.  

Let $X$ be the number of times the practioners are correct out of $210$ trials assuming the null hypothesis is true.  What distribution does $X$ have?

Assign the variable `X_dist` to the letter  (lower case  string) corresponding to the correct distribution for this scenario. 


 - (a): $X \sim Bin(210, 0.44)$
 - (b): $X \sim Bin(210, 0.5)$
 - (c): $X \sim Ber(.5)$
 - (d): $X \sim Ber(.44)$
 - (e): $X \sim Exp(.44)$
 - (f): $X \sim Exp(.55)$
 - (g): $X \sim Bin(100, 0.5)$

```python
X_dist = ...
```

```python
grader.check("q3g")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3h (2 pts) ## 

Create an array called `simulated_statistics` that contains 50,000 simulated values of the test statistic under the null hypothesis. Assume that the original sample consisted of 210 experiments. 

As usual, start by defining a function `one_simulated_statistic()` that simulates one value of the statistic. Your function may not use `np.random.DISTRIBUTION` where `DISTRIBUTION` is the distribution you chose in part 3g.  Your function should also use your `statistic` function from part 3e.  


We have included the code that plots the distribution of the simulated values. The red dot on the plot represents the observed statistic you found in Question 3f.


```python
def one_simulated_statistic():
    ...

num_simulations = 50000

simulated_statistics = ...


# Run the this cell a few times to see how the simulated statistic changes
one_simulated_statistic()
```

```python
# Run this cell to produce a histogram of the simulated statistics
plt.hist(simulated_statistics, density = True, ec= "white")
plt.xlabel('Simulated Statistic')
plt.ylabel('Percent per Unit')
plt.title('Histogram of Simulated Statistics')
plt.gca().yaxis.set_major_formatter(PercentFormatter(1))
plt.scatter(observed_statistic, -0.002, color='red', s=100);
plt.show()
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3i (1 pt) ## 

Using `simulated_statistics`, `observed_statistic`, and `num_simulations`, find the empirical p-value based on your simulation. 

```python
p_value = ...
p_value
```

```python
grader.check("q3i")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 3j (1 pt) ## 

What is the conclusion of this hypothesis test using the 5% cutoff for the p-value?

 - (a). The data leads us to reject Emily's null hypothesis that the practitioners were just guessing.  This indicates that something more than just random chance is at play. 

 - (b). Given the data, we fail to reject Emily's null hypothesis that the practitioners are answering independently and randomly (i.e. we fail to reject that the probability of any practitioner correctly answering is 50%).  This doesn't prove that her hypothesis is correct, it just says we didn't find strong enough evidence to reject it.

 - (c) We accept Emily's null hypothesis that the probability of any practitioner correctly answering is 50%.  This proves that she was correct.




```python
conclusion = ...
```

```python
grader.check("q3j")
```



Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!


***

[Back to top](#top)

<a id='p4'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4 ## 

Before beginning this section, please read [this section](https://inferentialthinking.com/chapters/11/2/Multiple_Categories.html#a-new-statistic-the-distance-between-two-distributions) of the textbook on TVD!

**Total variation distance (TVD)** is a special type of test statistic that we use when we want to compare two distributions of *categorical data*. It is often used when we observe that a set of observed proportions/probabilities is different than what we expect under the null model. 

Consider a six-sided die that we roll 6,000 times. If the die is fair, we would expect that each face comes up $\frac{1}{6}$ of the time. By random chance, a fair die won't always result in equal proportions (that is, we won't get exactly 1,000 of each face). However, if we suspect that the die might be unfair based on the data, we can conduct a hypothesis test using TVD to compare the expected [$\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$] distribution to what is actually observed.

In this part of the homework, we'll look at how we can use TVD to determine the effect that different factors have on happiness. 

We will be working with data from the [Gallup World Poll](https://www.gallup.com/analytics/349487/gallup-global-happiness-center.aspx#:~:text=World%20Happiness%20Report&text=Using%20the%20Gallup%20World%20Poll,about%20the%20World%20Happiness%20Report.) that is presented in the World Happiness Report, a survey of the state of global happiness. The survey ranked 155 countries by overall happiness and estimated the influence that economic production, social support, life expectancy, freedom, absence of corruption, and generosity had on population happiness. The study has been repeated for several years, but we'll be looking at data from the 2016 survey.

Run the cell below to load in the `happiness_scores` table.

```python
happiness_scores = pd.read_csv("data/happiness_scores.csv")
happiness_scores.head(5)
```

Participants in the study were asked to evaluate their life satisfaction from a scale of 0 (worst possible life) to 10 (best possible life). The responses for each country were averaged to create the `Happiness Score`.

The columns `Economy (GDP per Capita)`, `Family`, `Health (Life Expectancy)`, `Freedom`, `Trust (Government Corruption)`, and `Generosity` estimate the extent to which each factor influences happiness, both for better or for worse. The happiness score is the sum of these factors; the larger a factor is, the more it contributes to overall happiness. [In other words, if you add up all the factors (in addition to a "Difference from Dystopia" value we excluded in the dataset), you get the happiness score.]

Let's look at the different factors that affect happiness in the United States. Run the cell below to view the row in `us_happiness` that contains data for the United States.

```python
us_happiness = happiness_scores.loc[happiness_scores["Country"] == "United States"]
us_happiness
```

**To compare the different factors, we'll look at the proportion of the happiness score that is attributed to each variable. 
You can find these proportions in the table `us_happiness_factors` after running the cell below.**

*Note:* The factors shown in `us_happiness` don't add up exactly to the happiness score, so we adjusted the proportions to  only account for the data we have access to. The proportions were found by dividing each Happiness Factor value by the sum of all Happiness Factor values in `us_happiness`.

```python
us_happiness_factors = pd.read_csv("data/us_happiness_factors.csv")
us_happiness_factors
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4a (1 pt) ## 

Suppose we want to test whether or not each factor contributes the same amount to the overall Happiness Score. Define the null hypothesis, alternative hypothesis, and test statistic in the cell below. 

*Note:* Please format your answer as follows:
- Null Hypothesis: ...  
- Alternative Hypothesis: ...  
- Test Statistic: ...  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4b (1 pt) ## 

Write a function `calculate_tvd` that takes in the observed distribution (`obs_dist`) and expected distribution under the null hypothesis (`null_dist`) and calculates the total variation distance. Use this function to set `observed_tvd` to be equal to the observed test statistic. 


```python
null_distribution = np.ones(6) * (1/6)

def calculate_tvd(obs_dist, null_dist):
    ...
    
observed_tvd = ...
observed_tvd
```

```python
grader.check("q4b")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4c (2 pts) ## 

Create an array called `simulated_tvds` that contains 10,000 simulated values under the null hypothesis. Assume that the original sample consisted of 1,000 individuals. 

Then plot a density histogram of your simulated test statistics, as well as a red dot representing the observed value of the test statistic.  Include a title and label your x and y axes.


```python
...
# your code above this line
```

```python
grader.check("q4c")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4d (1 pt) ## 

Use your simulated statistics to calculate the empirical p-value of your test. Make sure that this number is consistent with what you observed in the histogram above.


```python
p_value_tvd = ...
p_value_tvd
```

```python
grader.check("q4d")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 4e (1 pt) ## 

What can you conclude about how each factor contributes to the overall happiness score in the US? Explain your answer using the results of your hypothesis test. Assume a significance level (i.e. p-value cutoff) of 5%. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

***

[Back to top](#top)

<a id='p5'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 5 (5 pts) ## 

#### Video Assignment

Complete this [Video Assignment: A/B Testing and Randomized Controlled Trials (click this link)](https://canvas.colorado.edu/courses/120673/assignments/2371674) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).

***

[Back to top](#top)

<a id='p6'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6: Deflategate ## 

On January 18, 2015, the Indianapolis Colts and the New England Patriots played the American Football Conference (AFC) championship game to determine which of those teams would play in the Super Bowl. After the game, there were allegations that the Patriots' footballs had not been inflated as much as the regulations required; they were softer. This could be an advantage, as softer balls might be easier to catch.

For several weeks, the world of American football was consumed by accusations, denials, theories, and suspicions: the press labeled the topic Deflategate, after the Watergate political scandal of the 1970's. The National Football League (NFL) commissioned an independent analysis. In this example, we will perform our own analysis of the data.

Pressure is often measured in pounds per square inch (psi). NFL rules stipulate that game balls must be inflated to have pressures in the range 12.5 psi and 13.5 psi. Each team plays with 12 balls. Teams have the responsibility of maintaining the pressure in their own footballs, but game officials inspect the balls. Before the start of the AFC game, all the Patriots' balls were at about 12.5 psi. Most of the Colts' balls were at about 13.0 psi. However, these pre-game data were not recorded.

During the second quarter, the Colts intercepted a Patriots ball. On the sidelines, they measured the pressure of the ball and determined that it was below the 12.5 psi threshold. Promptly, they informed officials. 

At half-time, all the game balls were collected for inspection. Two officials, Clete Blakeman and Dyrol Prioleau, measured the pressure in each of the balls. 

Here are the data. Each row corresponds to one football. Pressure is measured in psi. The Patriots ball that had been intercepted by the Colts was not inspected at half-time. Nor were most of the Colts' balls – the officials simply ran out of time and had to relinquish the balls for the start of second half play.

```python
football = pd.read_csv('data/deflategate.csv')
football
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6a (1 pt) ## 

For each of the 15 balls that were inspected, the two officials got different results. It is not uncommon that repeated measurements on the same object yield different results, especially when the measurements are performed by different people.

Add a new column to the  `football` dataframe called `Combined` where the entries are the average of the two measurements made for each ball.   Then drop the columns `Blakeman` and `Prioleau` from the `football` table.   


```python
...
# your code above this line

football
```

```python
grader.check("q6a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6b (2 pts) ## 

At a glance, it seems apparent that the Patriots' footballs were at a lower pressure than the Colts' balls. Because some deflation is normal during the course of a game, the independent analysts decided to calculate the **drop in pressure** from the start of the game. Recall that the Patriots' balls had all started out at about 12.5 psi, and the Colts' balls at about 13.0 psi. 

Complete the cell below so that `avg_drop_for_colts` and `avg_drop_for_patriots` evaluates to average drops in pressure for the Colts and Patriots respectively.

Hint: Start by creating a new column **PressureDrop** and assign it the drop in pressure for each ball compared to the pressure at the start of the game. 

```python
...
# Your code above this line

avg_drop_for_colts = ...
avg_drop_for_patriots = ...
```

```python
grader.check("q6b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6c (1 pt) ## 

The question now is why the Patriots' footballs had a larger drop in pressure, on average, than the Colts footballs. Could it be due to chance? 

How does chance come in here? Nothing was being selected at random. But we can make a chance model by hypothesizing that the 11 Patriots’ drops look like a random sample of 11 out of all the 15 drops, with the Colts’ drops being the remaining four.
That’s a completely specified chance model under which we can simulate data. So that's our null hypothesis.


Assign `alternative_statement_number` to the **number** next to the statement below that correctly describes the alternative hypothesis.  (For example, if you believe statement 1 is the correct alternative hypothesis, set `alternative_statement_number` = 1 in the code cell below).  

 1. There is no difference in the population distributions of pressure drops between all of the Patriots' balls and the Colts' balls.  Any difference observed in the random sample of balls measured at halftime was due to random chance.    

 2. The Patriots' pressure drops are too large, on average, to resemble a random sample drawn from all the drops. 
 
 3. The Patriots' pressure drops are too small, on average, to resemble a random sample drawn from all the drops.


```python
alternative_statement_number = ...
```

```python
grader.check("q6c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6d (1 pt) ## 

To test this hypothesis, let's follow our textbook's advice and choose a test statistic where the following statement is true: Large values of the test statistic favor the alternative hypothesis.

The correct test statistic for these hypotheses is one of the three options below. Which one is it? Assign the appropriate number to the variable `correct_test_stat`. 

1. "average pressure drop for Patriots - average pressure drop for Colts" 
2. "average pressure drop for Colts - average pressure drop for Patriots" 
3. "|average pressure drop for Patriots - average pressure drop for Colts|" 


```python
correct_test_stat = ...
correct_test_stat
```

```python
grader.check("q6d")
```

#### <hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6e (1 pt) ## 

Complete the cell below so that `observed_statistic_ab` evaluates to the observed value of the test statistic.  (The check below just checks that you have the correct format, not the correct actual value).

```python
observed_statistic_ab = ...

observed_statistic_ab
```

```python
grader.check("q6e")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6f (1 pt) ## 

To simulate our null hypothesis we will use a permutation test. Assign `permutation_test_reason` to the number corresponding to the reason why we use a permutation test in this situation.  Again, the built-in tests in this notebook only check if your answer is in the correct format.


1. You can always use a permutation test, regardless of your null hypothesis. 


2. Under the null hypothesis then it shouldn't matter which footballs are labeled Patriots and which are labeled Colts. The distributions of the two sets of pressure drops would be the same. We can simulate this by randomly shuffling the team **labels** in the `football` table.


3. Under the null hypothesis then it shouldn't matter which footballs are labeled Patriots and which are labeled Colts. The distributions of the two sets of pressure drops would be the same. We can simulate this by randomly shuffling the **rows** in the `football` table.



```python
permutation_test_reason = ...
```

```python
grader.check("q6f")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6g (2 pts) ## 

Create a new DataFrame `original_and_shuffled` that is a copy of the original `football` DataFrame, but that has an additional column **ShuffledLabel**  with shuffled Colts/Patriots labels.  

*Note:* Check out [12.1](https://inferentialthinking.com/chapters/12/1/AB_Testing.html#predicting-the-statistic-under-the-null-hypothesis) for a refresher on random permutations.


```python
original_and_shuffled = football.copy()

...
# Your code above this line

original_and_shuffled
```

```python
grader.check("q6g")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6h (4 pts) ## 

In the first cell below:

 - Define a function `simulate_one_statistic` that takes no arguments and returns one simulated value of the test statistic.  Refer to the code you have previously written in this problem, as you might be able to re-use some of it. 

In the 2nd cell below:

 - Complete the code to simulate 10,000 values of the statistic and store it in the array `simulated_statistics_ab`.

 - Then draw a density histogram with the empirical distribution of the statistic 
 
 - Include a red dot on your histogram at the value of `observed_statistic_ab`. 

 - Include a title for your histogram and label the x and y-axes.



```python
def simulate_one_statistic():
...

# Your code above this line


simulate_one_statistic()
```

```python
repetitions = 10000

simulated_statistics_ab = ...

...
# your code for histogram and observed statistic above this line

```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6i (1 pt) ## 

Use the simulation to find an empirical approximation to the p-value. Assign `p_val` to the appropriate p-value from this simulation. 

*Note:* Assume that we use the 5% significance cutoff for the p-value.


```python
p_val = ...

p_val
```

```python
grader.check("q6i")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## <span style='color:Red'>   Question 6j (2 pts) ## 

What is the conclusion of your test? Explain what this means in the context of this particular problem.    Can we make any causal conclusions from this test?  Why or why not?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In the actual case, the independent investigative team analyzed the data in several different ways, taking into account the laws of physics. The final report said,

“The average pressure drop of the Patriots game balls exceeded the average pressure drop of the Colts balls by 0.45 to 1.02 psi, depending on various possible assumptions regarding the gauges used, and assuming an initial pressure of 12.5 psi for the Patriots balls and 13.0 for the Colts balls.”  (Investigative report commissioned by the NFL regarding the AFC Championship game on January 18, 2015)

Our analysis shows an average pressure drop of about 0.73 psi, which is close to the center of the interval “0.45 to 1.02 psi” and therefore consistent with the official analysis.

But the all-important question in the football world was about causation: the question was whether the excess drop of pressure in the Patriots’ footballs was deliberate. If you are curious about the answer given by the investigators, here is the full report:
(https://www.wsj.com/public/resources/documents/Deflategate.pdf)


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 8!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 8 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW8 Autograded assignment in Gradescope.

* **Homework 8 Manually Graded**: Submit your hw08.PDF to the HW8 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw08.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2a (points: 1)
```
>>> assert get_hash(q2a_answer) == '6bb61e3b7bce0931da574d19d1d82c88'
```


Tests for q2b (points: 2)
```
>>> assert get_hash(round(q2b_answer, 3)) == '0a4a2fbfa050681b56c8bfce230ee97c'
```


Tests for q3a (points: 1)
```
>>> assert get_hash(null_hypothesis) == '92eb5ffee6ae2fec3ad71c777531578f'
```


Tests for q3b (points: 1)
```
>>> assert get_hash(alt_hypothesis) == '4a8a08f09d37b73795649038408b5f33'
```


Tests for q3c (points: 1)
```
>>> assert get_hash(np.round(expected_proportion_correct, 1)) == 'd310cb367d993fb6fb584b198a2fd72c'
```


Tests for q3d (points: 1)
```
>>> assert get_hash(valid_test_stat) == '8fa14cdd754f91cc6554c9e71929cce7'
```


Tests for q3e (points: 1)
```
>>> assert get_hash(int(round(statistic(0.5, 0.5) + statistic(0.4, 0.1), 1))) == '34173cb38f07f89ddbebc2ac9128303f'
```
```
>>> assert get_hash(int(statistic(0.4, 0.1) - statistic(0.1, 0.4))) == 'cfcd208495d565ef66e7dff9f98764da'
```


Tests for q3f (points: 1)
```
>>> assert get_hash(int(round(observed_statistic, 2))) == '1679091c5a880faf6fb5e6087eb1b2dc'
```


Tests for q3g (points: 1)
```
>>> assert get_hash(X_dist) == '92eb5ffee6ae2fec3ad71c777531578f'
```


Tests for q3i (points: 1)
```
>>> assert 0.08 < p_value < 0.09
```


Tests for q3j (points: 1)
```
>>> assert get_hash(conclusion) == '92eb5ffee6ae2fec3ad71c777531578f'
```


Tests for q4b (points: 1)
```
>>> assert type(observed_tvd) in set([float, np.float32, np.float64])
```
```
>>> assert get_hash(round(observed_tvd, 4)) == '5b9a8f5cf4912f44a3b12f1f0ce981cd'
```


Tests for q4c (points: 0)
```
>>> assert len(simulated_tvds) == 10000
```
```
>>> assert all(simulated_tvds >= 0)
```
```
>>> assert len(np.unique(simulated_tvds)) != 1
```
```
>>> assert np.mean(simulated_tvds) < 0.1
```


Tests for q4d (points: 1)
```
>>> assert 0 <= p_value_tvd <= 0.0001
```


Tests for q6a (points: 1)
```
>>> assert get_hash(football.size) == '34173cb38f07f89ddbebc2ac9128303f'
```
```
>>> assert get_hash(round(football['Combined'].mean(), 4)) == '3febcf8143e39bae52a9af09bb435c66'
```


Tests for q6b (points: 2)
```
>>> assert get_hash(round(avg_drop_for_colts, 4)) == '72e740c49ca1bcf70fc5e89e012a09e8'
```
```
>>> assert get_hash(round(avg_drop_for_patriots, 4)) == '49883c406681181ba6024b49ecd0cc3b'
```


Tests for q6c (points: 1)
```
>>> assert get_hash(alternative_statement_number) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q6d (points: [0, 1])
```
>>> assert type(correct_test_stat) == int
```
```
>>> assert get_hash(correct_test_stat) == 'c4ca4238a0b923820dcc509a6f75849b'
```


Tests for q6e (points: 1)
```
>>> assert get_hash(round(observed_statistic_ab, 4)) == '501bfe47af75dbbbe82fc9e745d803b9'
```


Tests for q6f (points: [0, 1])
```
>>> assert type(permutation_test_reason) == int
```
```
>>> assert get_hash(permutation_test_reason) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q6g (points: 2)
```
>>> assert all((label in original_and_shuffled.columns for label in ['Team', 'PressureDrop', 'ShuffledLabel']))
```
```
>>> assert len(original_and_shuffled) == 15
```
```
>>> assert all(football['Team'] == original_and_shuffled['Team'])
```
```
>>> assert all(football['PressureDrop'] == original_and_shuffled['PressureDrop'])
```
```
>>> assert all(original_and_shuffled['PressureDrop'] != original_and_shuffled['ShuffledLabel'])
```
```
>>> assert all((x in original_and_shuffled['ShuffledLabel'].values for x in ['Colts', 'Patriots']))
```


Tests for q6i (points: [0, 1])
```
>>> assert type(p_val) in set([float, np.float32, np.float64])
```
```
>>> assert 0.001 < p_val < 0.005
```
